<table style="width: 100%;">
<tr style="background-color: transparent;">
<td width="100px"><img src="https://cs104williams.github.io/assets/cs104-logo.png" width="90px" style="text-align: center"/></td>
<td>
  <p style="margin-bottom: 0px; text-align: left; font-size: 18pt;"><strong>CSCI 104: Data Science and Computing for All</strong><br>
                Williams College<br>
                Fall 2024</p>
</td>
</tr>


# Lab 1: Notebooks and Expressions

<font color='#B1008E'>
    
##### Learning objectives
- Navigate Jupyter notebooks (like this one)
- Write and evaluate some basic *expressions* in Python, the computer language of the course
- Call *functions* to use code other people have written
- Break down Python code into smaller parts to understand it
</font>

Sections 1-6 are to help you get started writing Python code.  Sections 7-9 are questions exercising your knowledge of those topics.  There are a number of ungraded "Exercises" for you to complete.  Any section with a title like "Part 4.1" will be graded for correctness when you submit your work.

This lab covers parts of [Chapter 3](http://www.inferentialthinking.com/chapters/03/programming-in-python.html) of the online textbook. You should read the examples in the book, but not right now. Instead, let's get started!

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 1. Jupyter Notebooks
This webpage is called a Jupyter notebook. A notebook is a place to write programs and view their results, and also to write text.

### 1.1. Text Cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You don't need to learn Markdown, but you might want to.

After you edit a text cell, click the "run cell" button at the top that looks like ▶ or hold down `shift` + `return` to confirm any changes. (Try not to delete the instructions of the lab.)

#### Exercise

This next paragraph is in its own text cell.  Edit it to add your name, and then click the "run cell" ▶ button or hold down `shift` + `return`.

My name is YOUR-NAME-HERE!

### 1.2. Code Cells

Cells can also contain code in the Python programming language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶ or hold down the `shift` key and press `return` or `enter`.

Try running the next cell:

```python
print("Hello, World!")
```

The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

### 1.3. Writing Jupyter Notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a code cell.  You can change it to a text cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (⟳) button in the menu bar, and choosing "Markdown". Or change it back to code by choosing "Code" 

<img src="dropdown_menu.png"/>

#### Exercise

Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell!
    
Run your cell to verify that it works.

### 1.4. Errors

Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever any of us write code, we'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

**Note:** In the toolbar, there is the option to click `Run > Run All Cells`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="error.jpg"/>

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you that Python can not interpret your expression because it does not follow Python's syntax.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask a neighbor, staff member, or the internet for help.)

#### Exercise

Try to fix the code above so that you can run the cell and see the intended message instead of an error.

### 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in, the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `[...]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection.  Oftentimes, the problem can be fixed by selecting **Kernel -> Interrupt Kernel** from the menu, or pressing the stop button from the toolbar at the top of the window.  Fix the issue in your code and then run the cell again.

See our [troubleshooting page](https://www.cs.williams.edu/~cs104/docs/how-to-jupyter/troubleshooting.html) for fixing other problems related to "stuck" kernels.

### 1.6. Submitting Your Work
All assignments in the course will be distributed as notebooks like this one, and you will submit your work from the notebook. We will use a system called Otter that checks your work and helps you submit. At the top of each assignment, you'll see a cell like the one below. Always run it when you open up one of our notebooks.

```python
# Initialize CS104 library and Otter Grader library
from cs104 import *

import otter
grader = otter.Notebook("lab01.ipynb")
```

You will submit your finished solution by running the export command below to generate a zip file containing your solution that you will then upload to Gradescope. It's fine to submit multiple times.  We will only grade your final submission.  You can run the export command now to see what it does, but don't worry about uploading anything yet.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 2. Numbers

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers. The expression `3.2500` evaluates to the number 3.25. (Run the cell and see.)

```python
3.2500
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(2)
3
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

### 2.1. Arithmetic
The line in the next cell subtracts.  Its value is what you'd expect.  Run it.

```python
3.25 - 1.5
```

Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/expressions.html) describes all the arithmetic operators used in the course.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses to compute a more meaningful number.

```python
3+6*5-6*3**2*2**3/4*7
```

```python
6+(6*5-(6*3))**2*((2**3)/4*7)
```

In standard math notation, the first expression is

$$3 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$6 + (6 \times 5 - (6 \times 3))^2 \times (\frac{(2^3)}{4} \times 7).$$

#### Exercise

Write a Python expression in this next cell that's equal to $5 \times (3 \frac{10}{11}) - 50 \frac{1}{3} + 2^{.5 \times 22} - \frac{7}{33} + 6$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus six.  By "$3 \frac{10}{11}$" we mean $3+\frac{10}{11}$, not $3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.

*Hint:* The correct output should start with a familiar number.

```python
...
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 3. Names
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large ruminant with horns, hoofs, and a tail!"  Instead, we just say, "cow!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

#### Exercise

Try writing code that uses a name (like `eleven`) that hasn't been assigned to anything.  You'll see an error!

```python
...
```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 841.25
monthly_salary = 2 * semimonthly_salary

number_of_months_in_a_year = 12

yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.  Blank lines in cells are okay, and using them to divide code into chunks can aid in helping you understand what code does.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 841.25
b = 2 * a
c = 12
d = c * b
d
```

**However, names are very important for making your code *readable* to yourself and others.**  The cell above is shorter, but it's totally useless without an explanation of what it does.

Another **style** point--in this and subsequent notebooks we name our variables with underscores, e.g. `yearly_salary`. Other styles of programming might use styles like "camel case," e.g. `YearlySalary.` Later in your career, you might develop opinions about your own programming style, but for now try to use variable names with underscores to be consistent. 

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 4. Checking your Code (5 pts)



Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. Please don't change the contents of the test cells. 

Go ahead and attempt the next question. Running the cell directly after it will test whether you have assigned `seconds_in_a_decade` correctly. If you haven't, this test will tell you the correct answer. Resist the urge to just copy it, and instead try to adjust your expression.

#### Part 4.1 (5 pts)


Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name seconds_in_a_decade.

seconds_in_a_decade = ...

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("p4.1")
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 5. Comments
You may have noticed these lines in the cell in which you answered the previous question:

    # Change the next line 
    # so that it computes the number of seconds in a decade 
    # and assigns that number the name seconds_in_a_decade.
    
This is called a *comment*. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 6. Calling Functions (5 pts)


The most common way to combine or manipulate values in Python is by calling functions. Python comes with many built-in functions that perform common operations.

For example, the `abs` function takes a single number as its argument and returns the absolute value of that number. Run the next two cells and see if you understand the output.

```python
abs(5)
```

```python
abs(-5)
```

#### Part 6.1 (5 pts)

Chunhua is on the corner of 7th Avenue and 42nd Street in Midtown Manhattan, and she wants to know far she'd have to walk to get to Gramercy School on the corner of 10th Avenue and 34th Street.

She can't cut across blocks diagonally, since there are buildings in the way.  She has to walk along the sidewalks.  Using the map below, she sees she'd have to walk 3 avenues (long blocks) and 8 streets (short blocks).  In terms of the given numbers, she computed 3 as the difference between 7 and 10, *in absolute value*, and 8 similarly.  

Chunhua also knows that blocks in Manhattan are all about 80m by 274m (avenues are farther apart than streets).  So in total, she'd have to walk $(80 \times |42 - 34| + 274 \times |7 - 10|)$ meters to get to the park.

<img width="50%" src="map.jpg"/>

Fill in the line `num_avenues_away = ...` in the next cell so that the cell calculates the distance Chunhua must walk and gives it the name `manhattan_distance`.  Everything else has been filled in for you.  **Use the `abs` function.** Also, be sure to run the test cell afterward to test your code.


```python
# Here's the number of streets away:
num_streets_away = abs(42-34)

# Compute the number of avenues away in a similar way:
num_avenues_away = ...

street_length_m = 80
avenue_length_m = 274

# Now we compute the total distance Chunhua must walk.
manhattan_distance = street_length_m * num_streets_away + avenue_length_m * num_avenues_away

# We've included this line so that you see the distance you've computed 
# when you run this cell.  
# You don't need to change it, but you can if you want.
manhattan_distance
```

```python
grader.check("p6.1")
```

#### Multiple arguments
Some functions take multiple arguments, separated by commas. For example, the built-in `max` function returns the maximum argument passed to it.

```python
max(2, -3, 4, -5)
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 7. Interpreting Graphs (10 pts)



The following scatter plot shows the duration and number of characters in the lyrics for the songs on Taylor Swift's first ten albums.  Specifically, the horizontal position of a dot measures the length of a song, and the vertical position measures the total number of characters in that songs lyrics.  (In computer science, we use the term character to refer to any letter, digit, punctuation mark or space in a piece of text.)  

<img src="taylor_swift.png"/>

#### Part 7.1 (5 pts)


Rounded to the nearest minute, how long is the song with the largest number of characters?

```python
duration_q1 = ...
```

```python
grader.check("p7.1")
```

#### Part 7.2 (5 pts)


Rounded to the nearest hundred, how many characters are in the song with the fewest characters sung per minute?

```python
characters_q2 = ...
```

```python
grader.check("p7.2")
```

To discover more interesting facts from this plot, read [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) of the textbook.

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 8. Names and Assignment Statements (15 pts)



#### Part 8.1 (5 pts)


When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

2. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

3. It should be `2 + 2 = 4`.

```python
names_q1 = ...
```

```python
grader.check("p8.1")
```

#### Part 8.2 (5 pts)


When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. Python cannot interpret the name `plus`, as it has not been defined.

<!--
BEGIN QUESTION
name: q3_2
manual: False
points:
 - 0
 - 4
-->

```python
names_q2 = ...
```

```python
grader.check("p8.2")
```

#### Part 8.3 (5 pts)


When you run the following cell, Python will, yet again, produce another cryptic error message.

```python
x = print(5)
y = x + 2
```

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q3` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.

<!--
BEGIN QUESTION
name: q3_3
manual: false
points:
 - 0
 - 4
-->

```python
names_q3 = ...
```

```python
grader.check("p8.3")
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 9. Differences between Majors (20 pts)



Adapted from information found on the Williams website, the table below displays the average number of degree recipients in three majors for the years 2008-2012 and 2018-2021.

| Major                              | 2008-2012    | 2018-2021   |
|------------------------------------|--------------|-------------|
| Comparative Literature             |  5           | 11          |
| Psychology                         | 62           | 45          |
| Mathematics                        | 53           | 61          |


#### Part 9.1 (5 pts)


Suppose you want to find the **biggest** absolute difference between the numbers of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 5 from 11) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers.

```python
biggest_change = ...
biggest_change
```

```python
grader.check("p9.1")
```

#### Part 9.2 (5 pts)


Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Comparative Literature
2. Psychology
3. Mathematics

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major.

```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("p9.2")
```

#### Part 9.3 (5 pts)


For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2012}} * 100}$ 

Fill in the code below such that `complit_relative_change`, `psych_relative_change` and `math_relative_change` are assigned to the relative changes for their respective majors.

```python
complit_relative_change = (abs(...) / 5) * 100

psych_relative_change = ...
math_relative_change = ...
print("complit_relative_change is ", complit_relative_change)
print("psych_relative_change is ", psych_relative_change)
print("math_relative_change is ", math_relative_change)
```

```python
grader.check("p9.3")
```

#### Part 9.4 (5 pts)


Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Comparative Literature
2. Psychology
3. Mathematics

Choose the number that corresponds to the major with the biggest relative change. 

```python
# Assign biggest_rel_change_major to the number corresponding to the major with the biggest relative change.
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("p9.4")
```

<hr class="m-0" style="border: 3px solid #500082;"/>

# You're Done!
Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to 
the corresponding assignment. For Lab N, the assignment will be called "Lab N Autograder".

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.


If that is the case, you will find results similar to the following on the Gradescope site once the autograder completes.

<img src="gradescope.png">

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for p4.1 (points: 5)
```
>>> check('seconds_in_a_decade' in vars())
```
failure_message: It looks like you didn't give anything the name
seconds_in_a_decade. Maybe there's a typo, or maybe you 
just need to run the cell below Question 3.2 where you defined 
seconds_in_a_decade. Click that cell and then click the "run
cell" button in the menu bar above.

success_message: seconds_in_a_decade has been successfully defined.
```
>>> check(seconds_in_a_decade != ...)
```
failure_message: It looks like you didn't change the cell to define
seconds_in_a_decade appropriately.  It should be a number,
computed using Python's arithmetic.  For example, this is
almost right: 
  seconds_in_a_decade = 10*365*24*60*60

success_message: seconds_in_a_decade has been given a value.
```
>>> check_str('seconds_in_a_decade == 315532800', locals())
```
failure_message: seconds_in_a_decade does not hold the right answer yet.
success_message: check(seconds_in_a_decade == 315532800)


Tests for p6.1 (points: 5)
```
>>> check_str('num_avenues_away == 3', locals())
```
failure_message: check(num_avenues_away == 3)
success_message: check(num_avenues_away == 3)
```
>>> check_str('manhattan_distance == 1462', locals())
```
failure_message: check(manhattan_distance == 1462)
success_message: check(manhattan_distance == 1462)


Tests for p7.1 (points: 5)
```
>>> check_str('duration_q1 in [1,2,3,4,5,6]', locals())
```
failure_message: duration_q1 should be a number between 1 and 6.
success_message: check(duration_q1 in [1,2,3,4,5,6])
```
>>> check_str('duration_q1 == 4', locals())
```
failure_message: duration_q1 is the wrong number.
success_message: check(duration_q1 == 4)


Tests for p7.2 (points: 5)
```
>>> check_str('characters_q2 in np.arange(0,5000, 100)', locals())
```
failure_message: characters_q2 should be a number between 1 and 3.
success_message: check(characters_q2 in np.arange(0,5000, 100))
```
>>> check_str('characters_q2 == 1000', locals())
```
failure_message: characters_q2 is the wrong number.
success_message: check(characters_q2 == 1000)


Tests for p8.1 (points: 5)
```
>>> check_str('names_q1 in [1,2,3]', locals())
```
failure_message: Be sure to set names_q1 to either 1, 2, or 3.
success_message: check(names_q1 in [1,2,3])
```
>>> check_str('names_q1 == 2', locals())
```
failure_message: names_q1 is the wrong number.
success_message: check(names_q1 == 2)


Tests for p8.2 (points: 5)
```
>>> check_str('names_q2 in [1,2,3,4]', locals())
```
failure_message: Be sure to set names_q2 to either 1, 2, 3, or 4.
success_message: check(names_q2 in [1,2,3,4])
```
>>> check_str('names_q2 == 4', locals())
```
failure_message: names_q2 is the wrong number.
success_message: check(names_q2 == 4)


Tests for p8.3 (points: 5)
```
>>> check_str('names_q3 in [1,2,3]', locals())
```
failure_message: Be sure to set names_q1 to either 1, 2, or 3.
success_message: check(names_q3 in [1,2,3])
```
>>> check_str('names_q3 == 2', locals())
```
failure_message: names_q3 is the wrong number.
success_message: check(names_q3 == 2)


Tests for p9.1 (points: 5)
```
>>> check_str('0 <= biggest_change', locals())
```
failure_message: check(0 <= biggest_change)
success_message: check(0 <= biggest_change)
```
>>> check_str('biggest_change == 17', locals())
```
failure_message: check(biggest_change == 17)
success_message: check(biggest_change == 17)


Tests for p9.2 (points: 5)
```
>>> check_str('smallest_change_major in [1,2,3]', locals())
```
failure_message: Be sure to set smallest_change_major to either 1, 2, or 3.
success_message: check(smallest_change_major in [1,2,3])
```
>>> check_str('smallest_change_major == 1', locals())
```
failure_message: check(smallest_change_major == 1)
success_message: check(smallest_change_major == 1)


Tests for p9.3 (points: 5)
```
>>> check_str('0 <= complit_relative_change <= 200', locals())
```
failure_message: check(0 <= complit_relative_change <= 200)
success_message: check(0 <= complit_relative_change <= 200)
```
>>> check_str('0 <= psych_relative_change <= 200', locals())
```
failure_message: check(0 <= psych_relative_change <= 200)
success_message: check(0 <= psych_relative_change <= 200)
```
>>> check_str('0 <= math_relative_change <= 200', locals())
```
failure_message: check(0 <= math_relative_change <= 200)
success_message: check(0 <= math_relative_change <= 200)
```
>>> check_str('complit_relative_change == approx(120.0)', locals())
```
failure_message: check(complit_relative_change == approx(120.0))
success_message: check(complit_relative_change == approx(120.0))
```
>>> check_str('psych_relative_change == approx(27.419354838709676)', locals())
```
failure_message: check(psych_relative_change == approx(27.419354838709676))
success_message: check(psych_relative_change == approx(27.419354838709676))
```
>>> check_str('math_relative_change == approx(15.09433962264151)', locals())
```
failure_message: check(math_relative_change == approx(15.09433962264151))
success_message: check(math_relative_change == approx(15.09433962264151))


Tests for p9.4 (points: 5)
```
>>> check_str('biggest_rel_change_major in [1,2,3]', locals())
```
failure_message: check(biggest_rel_change_major in [1,2,3])
success_message: check(biggest_rel_change_major in [1,2,3])
```
>>> check_str('biggest_rel_change_major == 1', locals())
```
failure_message: check(biggest_rel_change_major == 1)
success_message: check(biggest_rel_change_major == 1)
